---
  title: "Template Chapter 1"
  description: "This is a template chapter."
---

## Develop understanding of statistical variability

```yaml
type: NormalExercise 
lang: r
xp: 100 
skills: 1
key: fdb8f5d02d   
```


A statistical question is a question that can be answered by collecting data. Almost all data has some variability. Variability means that repeated measurements of similar items give different numbers (the numbers vary).

A deterministic question is a question with one exact answer. The number never changes.

Example of a statistical question:
How much time do students at  my school spend on math homework each week?

Example of a deterministic question:
How many pennies are in a dime?


`@instructions`
- Instruction 1
- Instruction 2
- Instruction 3
- Instruction 4

`@hint`
- Here is the hint for this setup problem. 
- It should get students 50% of the way to the correct answer.
- So don't provide the answer, but don't just reiterate the instructions.
- Typically one hint per instruction is a sensible amount.

`@pre_exercise_code`

```{r}
# Load datasets and packages here.
```

`@sample_code`

```{r}
# Your
# sample
# code
# should
# be
# ideally
# 10 lines or less,
# with a max
# of 16 lines.
```

`@solution`

```{r}
# Answer goes here
# Make sure to match the comments with your sample code
# to help students see the differences from solution
# to given.
```

`@sct`

```{r}
# Update this to something more informative.
success_msg("Some praise! Then reinforce a learning objective from the exercise.")
```

---

## Insert exercise title here

```yaml
type: PureMultipleChoiceExercise 
xp: 50 
key: 5e51bfcb9d   
```


A statistical question is a question that can be answered by collecting data. Almost all data has some variability. Variability means that repeated measurements of similar items give different numbers (the numbers vary).

A deterministic question is a question with one exact answer. The number never changes.

Example of a statistical question:
How much time do students at  my school spend on math homework each week?

Example of a deterministic question:
How many pennies are in a dime?

Abby went to the store a bought a bag of jelly beans, and has not opened it yet. Which of the following is not a statistical question about jelly beans?


`@hint`
If there is only one possible correct answer to the question, it is deterministic.
If there are multiple possible correct answers, or the correct answer could change, then it is a statistical question.

`@possible_answers`
- What is a typical number of jelly beans in a bag?
 - How much does Abby's bag weigh?
 - Which color is the most frequent?
 - If Abby closed her eyes, shook up the bag, reached in a grabbed a jelly bean, how likely will it be red?
 - What is the typical weight of a jelly bean?
 - How are the jelly beans distributed according to color?
 - All of the above
 - None of the above

`@feedbacks`
Try again
